<div class="alert alert-danger">

# 名前欄

</div>

以下に名前と学籍番号(STUDENT_ID)の入力をして下さい。

例)
```python
FAMILY_NAME = "Yamaguchi"# 姓をローマ字で入力
FIRST_NAME = "Daisuki"   #  名をローマ字で入力
USER = "a11yamaguchi"    # 公式メールアドレスのユーザ名
STUDENT_ID = "99999" # 学籍番号
```

```python
FAMILY_NAME = ""  # 姓をローマ字で入力
FIRST_NAME = ""  #  名をローマ字で入力
USER = ""  # 公式メールアドレスのユーザ名
STUDENT_ID = ""  # 99999 等と省略なしで入力
```

<div class="alert alert-danger">

# 提出方法(高専の皆さんへ)

</div>

以下のように**ディレクトリ(フォルダ)名を変更して提出**しましょう。
```python
元のディレクトリ名-学籍番号-名字(ローマ字)
```
例えば，元のディレクトリ名が`01-basic/`, 学籍番号が`99999`,名字が`Yamaguchi`の場合は以下になります。
```python
01-basic-99999-Yamaguchi/
```

<div class="alert alert-danger">

# 注意

</div>


- 以下の説明文を読み， 内容を納得できるように自分で試しなさい。
- **「# このセルで試してみよう」**, **「# このセルで動作確認をしよう」**と書かれているセルは自由に使って構いません。
- 必要に応じてセルを増やしても構いません。役立ちそうな情報もメモしておくと，今後プログラムを組む上で便利です。
- 表題等には適宜教科書のページを書いてあります。詳しい説明は教科書を参照しなさい。

<div class="alert alert-danger">

## 設問と解答用セル
    
</div>   

"**Q.**" から始まる設問に対する解答は必ず入力しなさい。解答用セルには以下のように表示される。

```python
# YOUR CODE HERE
raise NotImplementedError()
```

**注意**

- <span class="mark">raise NotImplementedError()の一行は**必ず消去**して</span>，<span class="mark"># YOUR CODE HERE の下に解答</span>を書くこと。
- <span class="mark">\# YOUR CODE HERE の一行は**消さない**</span>こと。どこに解答を書くべきかがわからなくなることがあります。
- なんらかの演算をしたり，命令文を作ったときは，その**結果が予想通りか否かを必ず確認しなさい**

<div class="alert alert-danger">

## 評価用セル
    
</div>

- **「評価用セル」と書かれているセルは採点のためのセル**です。
- **評価用コードが一部表示されている場合，それを実行すると課題の実行結果の確認**をできます。
- **表示されていない評価用コードも**あります。

<div class="alert alert-danger">

## 課題提出前の注意
    
</div>   

**[重要]** 以下の手順で，全ての課題が正常に動作するか確認してから提出しましょう。

1.  メニューバーの"Kernel"をクリック後，"Restart"を選択
2. メニューバーの"Cell"を，"Run All"を選択

```python
# このセルは自動評価を行うための命令です。内容は無視して，実行(Shift-Enter)のみをしてください。
#%%writefile unittests.py
#%load unittests.py
%run unittests.py
```

<div class="alert alert-success">

# 様々なパラメータ値に対する強化学習のパフォーマンス変化をグラフにする

</div>

## 準備

パラメータ$\epsilon$の様々な値に対する`mean_performance`の返り値を配列化して保存すれば，まとめて1つのグラフに表示できる。

まずは，課題シート`griwrorld2`で作成した以下のクラスや関数を入力しましょう。
- `class Qlearner`
- `play_env()`
- `mean_performance()` : 
- その他，上記のクラスや関数の実行に必要なもの

```python
# モジュールの読み込み
import numpy as np
import random as ra
import matplotlib.pyplot as plt
import gym
```

<!-- grade_id: A-Qlearner, points: 0, grade: false, solution: true -->
```python
# このセルに必要なクラスや関数を入力する
# YOUR CODE HERE
raise NotImplementedError()
```

## 学習パフォーマンスのグラフ作成

Q. 以下に，様々な$\epsilon$の値に対する`mean_performance()`の返り値`score_mean`を配列化するプログラム例を示す。
`mean_performance()`の中では`play_env()`を何度も実行してその結果を配列化しているのが，それと同様の処理内容である。
命令文を完成して，以下のようなグラフを作成しなさい。

![](images/RL-mean_performance.png)


**注意点**
- $\epsilon=0.2, 0.4, 0.6, 0.8$の場合について，`mean_performance()`の返り値`score_mean`をグラフ表示すること
- **軸名**は横軸`episodes`, 縦軸`score`とすること
- `reward=0`を示す**水平線**も描画すること
- **座標コンテナの名前**は`ax`にすること
- グラフは**折れ線グラフ**とする
- **凡例**(以下のグラフ右下部にある$\epsilon=0.0$, $\epsilon=0.2$等の一覧)をつけること。$\epsilon$の表示には`$\epsilon=$`と$\LaTeX$表記を用いること。
- グラフ化するデータの**エピソード数は100**にすること
- **フォントは見やすいサイズ**に調整すること

**プログラム例**

```python
env = gym.make("MDPGridworld-v1")
action_list = list(range(env.action_space.n))
agent = Qlearner(action_list)

e_range =   # epsilonの値設定。(上記仕様を満たす数値を要素とするndarrayを定義する)
max_episode = 100 # グラフ化するデータのエピソード数は100にする
mean =      # 0行max_episode列の2次元ndarrayを用意する。要素の初期化はしない。

for i,e in enumerate(e_range):
    agent=Qlearner(action_list, epsilon=e)
    """
    1. エージェントのインスタンスを生成 (例: agent = Qlearner(action_list, epsilon=e))
    2. mean_performance()を呼び出して，step_mean, score_mean, step_sd, score_sdを取得する
    3. for文の前で定義した配列meanに，新たな行としてscore_meanを追加           
    """

# グラフ作成
fig = plt.figure(figsize=(6, 4))
ax = fig.add_subplot()            # 座標コンテナ生成

for i in range(len(e_range)):
    # ここからは，ここまでの例を参考にグラフ作成
```


<!-- grade_id: A-RL-profile, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-RL-profile-title, points: 1, grade: true, solution: false -->
```python
# 評価用セル(軸ラベル確認)

_check_axes_title(ax,title=["episodes","score"])
print("ok!")
```

<!-- grade_id: E-RL-profile-data, points: 1, grade: true, solution: false -->
```python
# 評価用セルその2(データ確認)

ans_e_range = np.arange(0.2, 1.0, 0.2)

nt.assert_array_equal(e_range, ans_e_range, err_msg="εの値設定がおかしいようです。")
eq_(mean.shape, (len(ans_e_range), 100), msg="meanに格納されているデータ数が変です。")

pc=ut_msg(LinePlotChecker,ax,msg="座標コンテナが正常に定義されていないようです。")

eq_(len(pc.y_data), len(ans_e_range) + 1, msg="グラフ化しているデータ数が変です。")
eq_(len(pc.y_data[1]), 100, msg="グラフデータのエピソード数が100になっていないようです。")

```

<!-- grade_id: E-RL-profile-label, points: 1, grade: true, solution: false -->
```python
# 評価用セルその3(凡例(ラベル)確認)
_labels=[ "$\\epsilon=$"+str(np.round(i,2)) for i in ans_e_range ]
_check_labels(ax,_labels)
```

<div class="alert alert-success">

# 強化学習の学習後のパフォーマンスの分布をグラフにする

</div>

同じパラメータ値でも，シミュレーションを行うたびに最終エピソードにおける獲得スコアは異なる。
この最終獲得スコアの分布を，さまざまなパラメータ値に対して取得してグラフ化する方法を説明する。


## 準備

**Q.** 以下のような，`mean_performance()`の修正版`mean_performance2()`を作りなさい。

- **引数:**
    - 関数`mean_performance()`と同じ
- **返り値:** 
    - `play_env()`を指定回数(`n_trial`)を実行したとき，それぞれにおける**最終エピソードのステップ数と獲得得点**のリスト
    - 関数`mean_performance()`では，`play_env()`を指定回数(`n_trial`)を実行し，その実行結果として、各エピソードにおけるステップ数と獲得得点を格納する2つの配列(`n_trial`行`max_episode`列)を取得しているはず。**この配列それぞれの最後の列(最終エピソード)のデータ**(`n_trial`行1列)を**返り値(`step`と`score`)**とする。(平均値や標準偏差は求めないで返り値にする)

<!-- grade_id: A-mean_performance2, points: 0, grade: false, solution: true -->
```python
# ここで作ってみよう
# YOUR CODE HERE
raise NotImplementedError()
```

```python
# このセルで動作確認をしよう
```

<!-- grade_id: E-mean_performance2, points: 3, grade: true, solution: false -->
```python
# 評価用セル
class _Anonymous:
    def __init__(self, action_list, action="random", epsilon=0.0, gamma=0.0):
        self.action_list = action_list
        self.epsilon = epsilon
        self.gamma = gamma
        
        self.e_hist=[]
        self.r_hist=[]
        
        self.a = action if action in action_list else "random"

    def get_action(self, s):
        return ra.choice(self.action_list) if self.a == "random" else self.a

    
    def reset(self):
        self.e_hist+=[self.epsilon]
        self.r_hist+=[self.gamma]
        
        return 0

    def learning(self, s_old, action, r, s, terminal):
        return 0


def assert_mean_performance2(n_trial):
    action_list = list(range(env.action_space.n))
    step, score = mean_performance2(
        env=env, agent=_Anonymous(action_list, 0), n_trial=n_trial
    )
    nt.assert_array_equal(step.shape, (n_trial,), err_msg="step が n_trial行1列 ではありません。")
    nt.assert_array_equal(step.shape, (n_trial,), err_msg="score が n_trial行1列 ではありません。")

    nt.assert_equal(step, np.full(n_trial, 200.0), err_msg="stepの値が適切ではありません。")
    nt.assert_equal(score, np.full(n_trial, -600.0), err_msg="scoreの値が適切ではありません。")


assert_mean_performance2(5)
assert_mean_performance2(10)

print("passed!")
```

**Q.** いろいろな$\epsilon$の値に対して`mean_performance2()`を実行し，その学習結果(`step`と`score`)のうち`score`を以下のようにリスト化して返り値とする関数`get_data()`を作りなさい。


**引数**
- `e_range`: 学習結果を取得する$\epsilon$のリスト
- `mean_performance2()`と同じ引数。いずれも`mean_performance2()`を呼び出すときの引数とする。

**プログラム概要**

```python
def get_data(env, agent, e_range, n_trial=10, max_episode=100, learning=True):
    data =     # 0行n_trial列の2次元配列(ndarray)を用意する。

    for i,e in enumerate(e_range):
        """
        1. agent のインスタンス変数`epsilon`に`e`を代入する。
        2. mean_performance2()を呼び出して，返り値step, scoreを取得する
        3. 返り値scoreを1行n_trial列に変形(転置)して，配列dataに追加
        """
    return data
```

**実行例：** 以下の解答欄の次にあります。

<!-- grade_id: A-get_data, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

**Q.** `get_data()`を以下のように実行し，正常に動作するか確認しなさい。

```python
env = gym.make("MDPGridworld-v1")
action_list = list(range(env.action_space.n))
agent = Qlearner(action_list)

e_range =  # 0.2以上1.0未満の範囲で，0.2刻みの数値を要素とするndarrayを定義する
data=get_data(env,agent,e_range)
```

**データのグラフ化**は以下の解答欄の次でします。

<!-- grade_id: A-get_data_exec, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-get_data, points: 3, grade: true, solution: false -->
```python
_env = gym.make("MDPGridworld-v1")
_action_list = list(range(env.action_space.n))


def check_get_data(e_range, n_trial, learning=False):
    anon = _Anonymous(_action_list, 0)
    data = get_data(
        env=_env,
        agent=anon,
        e_range=e_range,
        n_trial=n_trial,
        max_episode=5,
        learning=learning,
    )
    nt.assert_array_equal(
        data.shape,
        (len(e_range), n_trial),
        err_msg="data が len(e_range)行 n_trial列 になってませんよ。",
    )

    # check if each epsilon is set to agent
    nt.assert_array_equal(
        np.repeat(e_range,n_trial), anon.e_hist, err_msg="e_histの各要素がエージェントのインスタンス変数epsilonに反映されてません。"
    )

    # check data
    ans_dat=np.full((len(e_range),n_trial),-600)
    nt.assert_array_equal(ans_dat,data,err_msg="dataの要素値が変です。")


check_get_data(e_range=np.arange(0.2, 1.0, 0.2), n_trial=1)
check_get_data(e_range=np.arange(0.2, 1.0, 0.2), n_trial=3)
print("good!")
```

## 折れ線グラフ

**Q.** 前問で作成したプログラムにより取得したデータ`data`から各$\epsilon$に対するデータの平均値$\pm$SDを計算し，以下の例のようなエラーバー付き折れ線グラフにしなさい。

![](images/RL-errorbar.png)

**注意点**
- 図のように，**$\epsilon$ は0.2以上1.0未満の範囲の0.2刻みの数値**とする
- 各軸ラベルをつける
- データ点(`marker`)は見やすいものを指定
- 横軸目盛りの数字は，データ点の$\epsilon$の値に対応させる
- `score=0`の水平線も描画
- フォントの大きさは見やすいサイズに調整

<!-- grade_id: A-errorbar, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-errorbar, points: 2, grade: true, solution: false -->
```python
# 評価用セル(軸タイトル)
_check_axes_title(ax, title=["$\epsilon$", "score"])
## check label
pc = LinePlotChecker(ax)

# 軸目盛り
ans_e_range = np.arange(0.2, 1.0, 0.2)
ut_msg(pc.assert_xticks_equal, ans_e_range, msg="横軸の目盛り位置が違うよ!")

# etc
ut_msg(pc.assert_num_lines, 4, msg="ラインの数が違うよ!?")

print("ok!")
```

## 棒グラフ

**Q.** 取得データ`data`の各$\epsilon$のデータ値の平均値$\pm$SDを，以下のサンプル例のようなエラーバー付き棒グラフにしなさい。

![](images/RL-bar.png)

**注意点**

- ここでも，**$\epsilon$ は0.2以上1.0未満の範囲の0.2刻みの数値**とする
- 各データの箱は`x=0,1,2,...`の位置に配置し，その位置に目盛りラベル`0.2, 0.4, 0.6, ..`を表示すると簡単。(横軸データとして`e_range`を第一引数に指定する方法もあるが，棒の幅指定等の設定も必要になるので面倒になる)
- 軸目盛りの数値を，$\epsilon$と$\gamma$の値に変更
    - リスト`e_range`の各値を目盛り数値にするには，`ax.bar()`の引数で`tick_label=e_range`と指定する
    - リスト要素(数値)を目盛りに使うために，**四捨五入**したいときには`np.round(list,decimals=0)`を使う。`decimals`は小数点以下第何桁を四捨五入するかを指定する引数

<!-- grade_id: A-bar, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-bar-title, points: 1, grade: true, solution: false -->
```python
# 評価用セル(軸タイトル)

_check_axes_title(ax, title=["$\epsilon$", "score"])

# 軸目盛り
pc = PlotChecker(ax)
ut_msg(pc.assert_xticks_equal, np.arange(4), msg="横軸の目盛りの位置が違うよ!")
ut_msg(
    pc.assert_xticklabels_equal,
    np.round(ans_e_range, decimals=2).astype("|U3"),
    msg="横軸の目盛り表示が違うよ!",
)
print("ok!")
```

<!-- grade_id: E-bar-num, points: 1, grade: true, solution: false -->
```python
# 評価用セル(そのほか)
pc = BarPlotChecker(ax)
ut_msg(pc.assert_num_bars,4,msg="バーの数が違うよ!")

print("good!")
```

## 箱ひげ図

**Q.** 取得したデータ`data`を，以下のサンプル例のような箱ひげ図にしなさい。

![](images/RL-boxplot.png)

**注意点**

- `ax.boxplot()`では，各データの箱は`x=1,2,...`の位置に配置される。
- 横軸目盛りのラベルを`e_range`の各値に変更するには，`ax.boxplot()`の引数で`labels=e_range`と指定する。(必要に応じて要素値を四捨五入)

<!-- grade_id: A-boxplot, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-boxplot, points: 2, grade: true, solution: false -->
```python
# 評価用セル(ラベル確認)

_check_axes_title(ax, ["$\epsilon$", "score"])
pc = PlotChecker(ax)
ut_msg(pc.assert_xticklabels_equal, ans_e_range.astype("|U3"), msg="目盛りのラベルが違いますよ。")

print("OK!")
```

<div class="alert alert-success">

# 複数のパラメータ値に対するパフォーマンスのグラフ化

</div>

## 準備

**Q**. 
以下は，すでに作成した`get_data()`を修正して，$\epsilon$だけでなく$\gamma$の値もいろいろ変えた場合の結果を取得するできるようにした関数`get_data2()`である。コメントに従ってこれを完成しなさい。

```python
def get_data2(env, agent, e_range, r_range, n_trial=10, max_episode=100, learning=True):
    mean =  # 行数が`r_range`の要素数，列数が`e_range`の要素数の2次元ndarrayを用意。初期化なしでOK
    sd =    # 同上
    
    for i, r in enumerate(r_range):
        # ここで agent のインスタンス変数`gamma`に`r`を代入。
        
        for j, e in enumerate(e_range):
            """
            1. agent のインスタンス変数`epsilon`に`e`を代入
            2. mean_performance2()を呼び出して，返り値step, scoreを取得
            3. mean[i,j]=  # 返り値scoreに格納されている要素値の平均値を求めて代入
            4. sd[i,j]=  # 返り値scoreに格納されている要素値の標準偏差を求めて代入
            """
    return mean,sd
```

**実行例：** 以下の解答欄の次にあります。

<!-- grade_id: A-get_data2, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

**Q.** `get_data2()`を以下のように実行し，正常に動作するか確認しなさい。

```python
env = gym.make("MDPGridworld-v1")
action_list = list(range(env.action_space.n))
agent = Qlearner(action_list)

e_range =   # epsilonの値設定。0.2以上1未満の範囲で，0.2刻みの数値を要素とするndarrayを定義
r_range =   # gammaの値設定。 0.2以1未満の範囲で，0.3刻みの数値を要素とするndarrayを定義

mean,sd=get_data2(env=env, agent=agent, e_range=e_range, r_range=r_range)
```

**データのグラフ化**は以下の解答欄の次でします。

<!-- grade_id: A-get_data2_exec, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-get_data2, points: 3, grade: true, solution: false -->
```python
_env = gym.make("MDPGridworld-v1")
_action_list = list(range(env.action_space.n))


def check_get_data2(e_range, r_range, n_trial, learning=False):
    anon = _Anonymous(_action_list, 0)
    mean,sd = get_data2(
        env=_env,
        agent=anon,
        e_range=e_range,
        r_range=r_range,
        n_trial=n_trial,
        max_episode=5,
        learning=learning,
    )
    nt.assert_array_equal(
        mean.shape,
        (len(r_range), len(e_range)),
        err_msg="mean が len(r_range)行 len(e_range)列 になってませんよ。",
    )
    nt.assert_array_equal(
        sd.shape,
        (len(r_range), len(e_range)),
        err_msg="sd が len(r_range)行 len(e_range)列 になってませんよ。",
    )

    # check if each epsilon/gamma is set to agent
    ans_e_hist=np.tile(np.repeat(e_range,n_trial),len(r_range))
    nt.assert_array_equal(
        ans_e_hist, anon.e_hist, err_msg="e_histの各要素がエージェントのインスタンス変数epsilonに反映されてません。"
    )
    
    ans_r_hist=np.repeat(r_range,n_trial*len(e_range))
    nt.assert_array_equal(ans_r_hist,anon.r_hist, err_msg="r_histの各要素がエージェントのインスタンス変数epsilonに反映されてません。")
    
    # check mean and sd
    ans_mean=np.full((len(r_range),len(e_range)),-600)
    ans_sd=np.full((len(r_range),len(e_range)),0)
    nt.assert_array_equal(ans_mean,mean,err_msg="meanの要素値が変です。")
    nt.assert_array_equal(ans_sd,sd,err_msg="sdの要素値が変です。")

_e_range=np.arange(0.2, 1.0, 0.2)
_r_range=np.arange(0.2, 1.0, 0.3)
check_get_data2(e_range=_e_range,r_range=_r_range, n_trial=1)
check_get_data2(e_range=_e_range,r_range=_r_range, n_trial=3)
_r_range=np.arange(0.2, 1.0, 0.2)
check_get_data2(e_range=_e_range,r_range=_r_range, n_trial=2)


print("good!")
```

## 折れ線グラフ

**Q.** 前問で取得した各$\epsilon$, $\gamma$に対するデータ`mean`, `sd`を用いて，以下の例のようなエラーバー付き折れ線グラフを作りなさいる

![](images/RL-errorbar2.png)

**注意点**

- データ点(`marker`)は見やすいものを指定
- `score=0`の水平線も描画
- 各軸ラベルをつける
- 横軸目盛りの数字は，データ点の$\epsilon$の値に対応させる
- 凡例をつける
- フォントは見やすいサイズに調整
- 課題シート`matplotlib2`で説明したように，`plot()`と`errorbar()`を併用する。
    - 同じデータに対する`plot()`と`errorbar()`の色を揃える必要があるが，色指定には`tab10`以外のカラーマップ(下記参照)を利用する。

**カラーマップについて**

matplotlibのデフォルトのカラーマップ`tab10`を利用する場合，以下を命令文に加えると，色の番号指定が可能になる。
```python
cmap = plt.get_cmap("tab10")
...
for i in range(...):
    ax.plot(..., color=cmap(i), ...)
    ax.errorbar(..., color=cmap(i), ...)
```
- 次のセルでは，**`tab10`以外のカラーマップを利用**すること
- **カラーマップ**は好きなものに変更
    - [カラーマップ一覧](https://matplotlib.org/tutorials/colors/colormaps.html?highlight=colormap)の下の方にある`Qualitative colormaps`から選択すると良い

<!-- grade_id: A-errorbar2, points: 0, grade: false, solution: true -->
```python
# 解答欄

fig = plt.figure()
cmap = plt.get_cmap("tab10") # カラーマップを利用。tab10以外のものを指定してください。

ax = fig.add_subplot(111)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-errorbar2_title, points: 1, grade: true, solution: false -->
```python
# 評価用セルその1: 軸タイトルおよび凡例確認

_check_axes_title(ax, title=["$\epsilon$", "score"])

pc = LinePlotChecker(ax)
pc.find_permutation("labels", ["$\gamma=$0.2", "$\gamma=$0.5", "$\gamma=$0.8"])
print("ok!")
```

<!-- grade_id: E-errorbar2_xticks, points: 1, grade: true, solution: false -->
```python
# 評価用セルその2: 軸目盛り

```

<!-- grade_id: E-errorbar2_data, points: 1, grade: true, solution: false -->
```python
# 評価用セルその3: データ確認

pc=ut_msg(LinePlotChecker,ax,msg="座標コンテナが正常に定義されていないようです。")

_assert_array_involve(pc.x_data, ans_e_range, msg="グラフ化されているデータ数が変です")
_assert_array_involve(pc.x_data, [0.0, 1.0], msg="score=0の線は描画されてますか?")
_assert_array_involve(pc.y_data, [0.0, 0.0], msg="score=0の線は描画されてますか?")
```

<!-- grade_id: E-errorbar2_colormap, points: 1, grade: true, solution: false -->
```python
# 評価用セルその4: カラーマップ確認

cmap_tab10 = plt.get_cmap("tab10")
try:
    assert not all(np.array_equal(cmap(i),cmap_tab10(i)) for i in range(8))
except:
    _alert("tab10以外のカラーマップを使いましょう")
    raise ValueError
```

## ヒートマップ

パラメータ$\epsilon$, $\gamma$の値のさまざまな組み合わせに対して得た結果一覧(`mean`)を，課題シート`matplotlib`で作成したプログラムを用いて**ヒートマップ**を作ると，以下の図が得られる。

![](images/RL-heatmap.png)

<!-- grade_id: cell-bf0c74f44751b179, points: 0, grade: false, solution: true -->
```python
# このセルで試してみよう

# YOUR CODE HERE
raise NotImplementedError()
```

**Q.** ヒートマップを以下の例のように修正しなさい。

![](images/RL-heatmap2.png)

**注意点**

- **各軸ラベル**をつける(`$\epailon$`と`$\gamma$`)
- **軸目盛り位置**を，ヒートマップの各セルの中心位置 ($x=0.5,1.5,...$の位置)にする
    - 横軸は`ax.set_xticks()`で目盛位置を指定
- **軸目盛りの数値**を，$\epsilon$と$\gamma$の値に変更
    - 目盛の数値はリストにして，`ax.set_xticklabels(labels=...)`と指定。(必要に応じて四捨五入)
- ヒートマップの各マス目の縦横比(**アスペクト比**)を1:1にし，マス目を正方形にする
    - `ax.set_aspect('equal')` を命令文に追加すればOK
- **フォントは見やすいサイズ**に調整
- **カラーマップは好きなもの**に変更
    - [カラーマップ一覧](https://matplotlib.org/tutorials/colors/colormaps.html?highlight=colormap)の下の方の一覧から選択すると良い。ただし**配色数が限定される Qualitative colormaps 以外**を選択すること

<!-- grade_id: A-RL-heatmap, points: 0, grade: false, solution: true -->
```python
# 解答欄
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: E-RL-heatmap-title, points: 1, grade: true, solution: false -->
```python
# 評価用セルその1: 軸タイトル
_check_axes_title(ax, title=["$\epsilon$", "$\gamma$"])
```

<!-- grade_id: E-RL-heatmap-ticks, points: 1, grade: true, solution: false -->
```python
# 評価用セルその2: 軸目盛り

ans_xticks=np.arange(len(e_range)) + 0.5
ans_yticks=np.arange(len(r_range)) + 0.5

pc = PlotChecker(ax)
try:
    pc.assert_xticks_equal(ans_xticks)
    pc.assert_yticks_equal(ans_yticks)
    pc.assert_xticklabels_equal(ans_e_range.astype("|U3"))
    pc.assert_yticklabels_equal(ans_r_range.astype("|U3"))
except:
    display("x軸かy軸の目盛りか，目盛りラベルが変なようです。")
```

<!-- grade_id: E-RL-heatmap-aspect, points: 1, grade: true, solution: false -->
```python
# 評価用セルその3: Aspect Ratio
eq_(ax.get_aspect(),'equal',msg="アスペクト比の設定をしましょう。")
```

```python
# このセルは，自動採点の際に以降のセルを実行しないための命令です。無視してください。
raise SystemExit("Stop right here!")
```

<div class="alert alert-success">

# 以下には，自由にセルを作成して，いろいろと動作確認して構いません。

</div>

```python

```